# Activity: Selective Subsets

## Introduction

In this activity you will practice selecting subsets of data from a DataFrame using Pandas.
This activity will cover the following topics:
- Creating masks
- Negating masks
- Masks with slicing
- Null value masks


#### Question 1

Create a `DataFrame` called `df` from the given CSV file `movie_data.csv`, and then create a mask called `before_millennium` to select all movies that were released before 2000.


<!-- grade_id: cell-240a9f6b83e2bc8e, points: 0, grade: false, solution: true -->
```python
import pandas as pd

# Your code here
df = pd.read_csv('C:/Users/Falcon/Desktop/Data Analyst Professional Training/Python Data Analytics_5 Weeks/movie_data.csv')
# your code here
```

Output:
```
Title  Year Released  Rating  Box Office ($M)
0   The Shawshank Redemption           1994     9.3             58.3
1              The Godfather           1972     9.2            246.1
2            The Dark Knight           2008     9.0           1005.0
3               Pulp Fiction           1994     8.9            213.9
4           Schindler's List           1993     8.9            321.3
..                       ...            ...     ...              ...
70                 Andhadhun           2018     8.3             48.0
71                 Gully Boy           2019     8.1             62.0
72            Dil Chahta Hai           2001     8.1             13.0
73          Dil To Pagal Hai           1997     7.1             11.0
74              Om Shanti Om           2007     6.7             23.0

[75 rows x 4 columns]
```

<!-- grade_id: cell-c67cccfe6b6864c5, points: 1, grade: true, solution: false -->
```python
# Question 1 Grading Checks

assert isinstance(df, pd.DataFrame), 'Did you create a DataFrame called df?'

```

```python
before_millenium = df[df['Year Released'] < 2000]
before_millenium
```

Output:
```
Title  Year Released  Rating  Box Office ($M)
0      The Shawshank Redemption           1994     9.3             58.3
1                 The Godfather           1972     9.2            246.1
3                  Pulp Fiction           1994     8.9            213.9
4              Schindler's List           1993     8.9            321.3
5                    Fight Club           1999     8.8            100.9
6                  Forrest Gump           1994     8.8            678.2
8                    The Matrix           1999     8.7            463.5
10     The Silence of the Lambs           1991     8.6            272.7
12               The Green Mile           1999     8.6            286.8
13       The Godfather: Part II           1974     9.0             48.0
16           The Usual Suspects           1995     8.5             34.4
17          Saving Private Ryan           1998     8.6            482.3
18                        Se7en           1995     8.6            327.3
21                
```

#### Question 2

Using the `before_millennium` mask from Question 1, assign the titles of every movie that was released after 2000 to a `Series` called `newer_titles`.


<!-- grade_id: cell-a008246edc6867e8, points: 0, grade: false, solution: true -->
```python
# Your code here
newer_titles = df.loc[df['Year Released'] > 2000, 'Title']
# your code here
newer_titles
```

Output:
```
2                                       The Dark Knight
7                                             Inception
9     The Lord of the Rings: The Fellowship of the Ring
14                The Lord of the Rings: The Two Towers
15                                         The Departed
19        The Lord of the Rings: The Return of the King
20                                          The Pianist
22                                               Avatar
23                                 Inglourious Basterds
24                                         The Revenant
25                                     Django Unchained
26                                   The Social Network
27                                         Interstellar
28                             The Grand Budapest Hotel
29                                              Dunkirk
30                                     A Beautiful Mind
33                                   The Shape of Water
34                                     12 Years 
```

<!-- grade_id: cell-60a9d4c511c889a5, points: 1, grade: true, solution: false -->
```python
# Question 2 Grading Checks

assert isinstance(newer_titles, pd.Series), 'Did you create a Series called newer_titles?'

```

```python
newer_titles_df = newer_titles.to_frame()
newer_titles_df
```

Output:
```
Title
2                                     The Dark Knight
7                                           Inception
9   The Lord of the Rings: The Fellowship of the Ring
14              The Lord of the Rings: The Two Towers
15                                       The Departed
19      The Lord of the Rings: The Return of the King
20                                        The Pianist
22                                             Avatar
23                               Inglourious Basterds
24                                       The Revenant
25                                   Django Unchained
26                                 The Social Network
27                                       Interstellar
28                           The Grand Budapest Hotel
29                                            Dunkirk
30                                   A Beautiful Mind
33                                 The Shape of Water
34                                   12 Years a Slave
35                    
```

#### Question 3

Create a mask to select movies with a `Rating` of `8.9` and a `Box Office ($M)` value higher than `1000.0`. Assign the resulting `Series` to a variable called `popular_pg_movies`.


<!-- grade_id: cell-61027e670582168c, points: 0, grade: false, solution: true -->
```python
# Your code here
popular_pg_movies = df.loc[(df['Rating'] == 8.9) & (df['Box Office ($M)'] > 1000.0), 'Title']
# your code here
popular_pg_movies
```

Output:
```
19    The Lord of the Rings: The Return of the King
Name: Title, dtype: object
```

<!-- grade_id: cell-e29d54428ab17eed, points: 1, grade: true, solution: false -->
```python
# Question 3 Grading Checks

assert isinstance(popular_pg_movies, pd.Series), 'Did you create a Series called popular_pg_movies?'

```

#### Question 4

Create a mask to select movies with a null value for `Box Office ($M)` or `Rating`. Assign the resulting `Series` to a variable called `missing_info`.


<!-- grade_id: cell-ea4f4905c6892e32, points: 0, grade: false, solution: true -->
```python
# Your code here
missing_info = df.loc[(df['Box Office ($M)'].isna()) | (df['Rating'].isna()), 'Title']
# your code here
missing_info

```

Output:
```
34    12 Years a Slave
Name: Title, dtype: object
```

<!-- grade_id: cell-543bf8c0aea1a245, points: 1, grade: true, solution: false -->
```python
# Question 4 Grading Checks

assert isinstance(missing_info, pd.Series), 'Did you create a Series called missing_info?'

```