# Activity: Removing Data

## Introduction

In this activity you will practice using Pandas functionality to check for and remove any unwanted data from a dataset.
This activity will cover the following topics:
- Removing columns from a DataFrame
- Removing rows from a DataFrame
- Removing rows based on a condition
- Checking for duplicate data


#### Question 1

Create a `DataFrame` called `df` from the given CSV file `exotic_plants_data.csv`, then drop the column `Type` and assign the result to a new `DataFrame` called `df_no_type`.


<!-- grade_id: cell-99d36e8c8c66ea51, points: 0, grade: false, solution: true -->
```python
import pandas as pd

df = pd.read_csv('C:/Users/Falcon/Desktop/Data Analyst Professional Training/Python Data Analytics_5 Weeks/exotic_plants_data.csv')
# Your code here
df
# your code here
df_no_type = df.drop('Type', axis=1)
df_no_type
```

Output:
```
Plant Name                     Origin  Height (cm)
0             Orchid                   Tropical           30
1               Fern                   Tropical           40
2             Bamboo                       Asia          600
3             Cactus                    America           60
4   Bird of Paradise                     Africa          150
..               ...                        ...          ...
71             Ficus                       Asia          200
72         Columbine              North America           30
73           Jasmine                       Asia           90
74           Fuchsia  Central and South America           40
75          Amaranth                    Various           80

[76 rows x 3 columns]
```

<!-- grade_id: cell-84571c0862a237ac, points: 1, grade: true, solution: false -->
```python
# Question 1 Grading Checks

assert isinstance(df, pd.DataFrame), 'Have you created a DataFrame named df?'
assert isinstance(df_no_type, pd.DataFrame), 'Have you created a DataFrame named df_no_type?'

```

#### Question 2

Check the `df` `DataFrame` for any duplicate rows and assign the result to a new `DataFrame` called `df_duplicates`.


<!-- grade_id: cell-ee673ddb1d8d4343, points: 0, grade: false, solution: true -->
```python
# Your code here
df_duplicates = df[df.duplicated()]
# your code here
df_duplicates
```

Output:
```
Plant Name         Type          Origin  Height (cm)
6             Cactus    Succulent         America           60
30         Rafflesia       Flower  Southeast Asia           20
47      Kangaroo Paw       Flower       Australia           60
48     Bougainvillea        Shrub   South America          400
49  Bird of Paradise   Ornamental          Africa          150
50     Venus Flytrap  Carnivorous   North America           15
51              Rose       Flower            Asia           60
```

<!-- grade_id: cell-496fa33508cec3f2, points: 1, grade: true, solution: false -->
```python
# Question 2 Grading Checks

assert isinstance(df_duplicates, pd.DataFrame), 'Have you created a DataFrame named df_duplicates?'

```

#### Question 3

Check the `df` `DataFrame` for any duplicate rows based on the `Plant Name` and `Type` columns and assign the result to a new `DataFrame` called `df_plant_type_duplicates`.


<!-- grade_id: cell-6f7859a5535a9c9b, points: 0, grade: false, solution: true -->
```python
# Your code here
df_plant_type_duplicates = df[df.duplicated(subset=['Plant Name', 'Type'], keep=False)]
# your code here
df_plant_type_duplicates.head()

```

Output:
```
Plant Name        Type         Origin  Height (cm)
2             Bamboo       Grass           Asia          600
3             Cactus   Succulent        America           60
4   Bird of Paradise  Ornamental         Africa          150
6             Cactus   Succulent        America           60
10     Bougainvillea       Shrub  South America          400
```

<!-- grade_id: cell-67e41e79158b9f13, points: 1, grade: true, solution: false -->
```python
# Question 3 Grading Checks

assert isinstance(df_plant_type_duplicates, pd.DataFrame), 'Have you created a DataFrame named df_duplicates?'

```

#### Question 4

Create a mask called `clean_mask` that will clean up any duplicates in the `df` DataFrame that have the same `Plant Name` and `Origin` and only keep the most up-to-date duplicate entry.


<!-- grade_id: cell-722a62389afc976a, points: 0, grade: false, solution: true -->
```python
# Your code here
clean_mask = df.duplicated(subset=['Plant Name','Origin'], keep= 'last')
# your code here
df_clean_mask = df[-clean_mask]
df_clean_mask

```

Output:
```
Plant Name          Type                     Origin  Height (cm)
0        Orchid    Ornamental                   Tropical           30
1          Fern  Ground Cover                   Tropical           40
5   Banana Tree          Tree                   Tropical          300
6        Cactus     Succulent                    America           60
7      Monstera    Ornamental                   Tropical           70
..          ...           ...                        ...          ...
71        Ficus          Tree                       Asia          200
72    Columbine        Flower              North America           30
73      Jasmine         Shrub                       Asia           90
74      Fuchsia        Flower  Central and South America           40
75     Amaranth        Flower                    Various           80

[66 rows x 4 columns]
```

<!-- grade_id: cell-a3bd7512e6b2d51c, points: 1, grade: true, solution: false -->
```python
# Question 4 Grading Checks

assert isinstance(clean_mask, pd.Series), 'Have you created a Series named clean_mask?'

```