# Activity: Basic Exploration

## Introduction

In this activity you will practice using exploration methods on a data set containing games of online chess*.
This activity includes some or all of the following, not necessarily in this order:
- Viewing the data
- Finding the mean
- Finding the median
- Standard deviation
- Aggregations
- Grouping

*The data set is from [Kaggle](https://www.kaggle.com/datasets/mysarahmadbhat/online-chess-games).


## Note

This data set is larger than those used in previous activities. Please run the cell below which uses the `info()` method to get a sense of the data before you begin.

```python
!pip install kaggle kagglehub
```

Output:
```
Collecting kaggle
  Downloading kaggle-1.7.4.2-py3-none-any.whl.metadata (16 kB)
Collecting kagglehub
  Downloading kagglehub-0.3.11-py3-none-any.whl.metadata (32 kB)
Requirement already satisfied: bleach in c:\users\falcon\onedrive\apps\anaconda3\lib\site-packages (from kaggle) (4.1.0)
Requirement already satisfied: certifi>=14.05.14 in c:\users\falcon\onedrive\apps\anaconda3\lib\site-packages (from kaggle) (2025.1.31)
Requirement already satisfied: charset-normalizer in c:\users\falcon\onedrive\apps\anaconda3\lib\site-packages (from kaggle) (2.0.4)
Requirement already satisfied: idna in c:\users\falcon\onedrive\apps\anaconda3\lib\site-packages (from kaggle) (3.7)
Requirement already satisfied: protobuf in c:\users\falcon\onedrive\apps\anaconda3\lib\site-packages (from kaggle) (3.20.3)
Requirement already satisfied: python-dateutil>=2.5.3 in c:\users\falcon\onedrive\apps\anaconda3\lib\site-packages (from kaggle) (2.9.0.post0)
Requirement already satisfied: python-slugify in c:\users\f
```

```python
import kagglehub

# Download latest version of the dataset
path = kagglehub.dataset_download("mysarahmadbhat/online-chess-games")

print("Path to dataset files:", path)
```

Output:
```
Downloading from https://www.kaggle.com/api/v1/datasets/download/mysarahmadbhat/online-chess-games?dataset_version_number=1...
100%|█████████████████████████████████████████████████████████████████████████████| 2.54M/2.54M [00:01<00:00, 2.28MB/s]
Extracting files...
Path to dataset files: C:\Users\Falcon\.cache\kagglehub\datasets\mysarahmadbhat\online-chess-games\versions\1
```

```python
import pandas as pd

df = pd.read_csv('chess_games.csv')
df.info()

```

Output:
```
<class 'pandas.core.frame.DataFrame'>
RangeIndex: 20058 entries, 0 to 20057
Data columns (total 17 columns):
 #   Column             Non-Null Count  Dtype 
---  ------             --------------  ----- 
 0   game_id            20058 non-null  int64 
 1   rated              20058 non-null  bool  
 2   turns              20058 non-null  int64 
 3   victory_status     20058 non-null  object
 4   winner             20058 non-null  object
 5   time_increment     20058 non-null  object
 6   white_id           20058 non-null  object
 7   white_rating       20058 non-null  int64 
 8   black_id           20058 non-null  object
 9   black_rating       20058 non-null  int64 
 10  moves              20058 non-null  object
 11  opening_code       20058 non-null  object
 12  opening_moves      20058 non-null  int64 
 13  opening_fullname   20058 non-null  object
 14  opening_shortname  20058 non-null  object
 15  opening_response   1207 non-null   object
 16  opening_variation  14398 non-null  objec
```

#### Question 1

Create two `DataFrame` objects called `first_three` and `last_three` and assign the first and last three rows of the data set to them, respectively.


<!-- grade_id: cell-254b741d2204b531, points: 0, grade: false, solution: true -->
```python
# Your code here
first_three = df.head(3)
last_three = df.tail(3)
# your code here
first_three
last_three
```

Output:
```
game_id  rated  turns victory_status winner time_increment  white_id  \
0        1  False     13    Out of Time  White           15+2  bourgris   
1        2   True     16         Resign  Black           5+10      a-00   
2        3   True     61           Mate  White           5+10    ischia   

   white_rating   black_id  black_rating  \
0          1500       a-00          1191   
1          1322  skinnerua          1261   
2          1496       a-00          1500   

                                               moves opening_code  \
0  d4 d5 c4 c6 cxd5 e6 dxe6 fxe6 Nf3 Bb4+ Nc3 Ba5...          D10   
1  d4 Nc6 e4 e5 f4 f6 dxe5 fxe5 fxe5 Nxe5 Qd4 Nc6...          B00   
2  e4 e5 d3 d6 Be3 c6 Be2 b5 Nd2 a5 a4 c5 axb5 Nc...          C20   

   opening_moves                        opening_fullname    opening_shortname  \
0              5        Slav Defense: Exchange Variation         Slav Defense   
1              4  Nimzowitsch Defense: Kennedy Variation  Nimzowitsch Defense   
2    
```

<!-- grade_id: cell-92f6542acea01bd8, points: 1, grade: true, solution: false -->
```python
# Question 1 Grading Checks

assert first_three.shape == (3, 17), 'Make sure that you chose only the first three rows.'
assert last_three.shape == (3, 17), 'Make sure that you chose only the last three rows.'

```

#### Question 2

Create two new `DataFrame` objects called `white_lower_rating` and `white_higher_rating` that are assigned the rows of data where the white player's rating is less than 1200 and greater than or equal to 1800, respectively.


<!-- grade_id: cell-c55abfb6732d83be, points: 0, grade: false, solution: true -->
```python
# Your code here
white_lower_rating = df[df['white_rating'] < 1200]
white_higher_rating = df[df['white_rating'] >= 1800]
# your code here
white_lower_rating
white_higher_rating
```

Output:
```
game_id  rated  turns victory_status winner time_increment  \
14          15  False     31           Mate  White          15+15   
15          16  False     43         Resign  Black          15+15   
16          17  False     52         Resign  Black          15+15   
17          18  False     66           Mate  Black          15+16   
18          19  False    101         Resign  Black          15+15   
...        ...    ...    ...            ...    ...            ...   
19840    19841  False     95         Resign  White           10+0   
19859    19860  False     46           Mate  Black           10+0   
19894    19895   True     54           Mate  Black           10+0   
19921    19922   True     54           Mate  Black          30+30   
20037    20038   True     32           Mate  Black           10+0   

                white_id  white_rating           black_id  black_rating  \
14     shivangithegenius          1094      sureka_akshat          1141   
15         sureka_akshat    
```

<!-- grade_id: cell-060c576da847aa8a, points: 1, grade: true, solution: false -->
```python
# Question 2 Grading Checks

assert isinstance(white_lower_rating, pd.DataFrame), 'Make sure that you are creating a DataFrame object called white_lower_rating.'
assert isinstance(white_higher_rating, pd.DataFrame), 'Make sure that you are creating a DataFrame object called white_higher_rating.'

```

#### Question 3

Using the `black_rating` column, create a `DataFrame` object called `top_10_percent_black` which is assigned the top 10% of black players by rating. That is, the only rows of where the `black_rating` is higher than 90% of all the `black_rating` values.

<!-- grade_id: cell-1c295e601c099a68, points: 0, grade: false, solution: true -->
```python
# Your code here
top_10_percent_black = df[df['black_rating'] >= df['black_rating'].quantile(0.9)]
# your code here
top_10_percent_black
```

Output:
```
game_id  rated  turns victory_status winner time_increment  \
7            8  False      9         Resign  Black          15+30   
78          79   True     47         Resign  Black           15+5   
203        204   True     28         Resign  Black            7+2   
291        292   True     69           Draw   Draw           20+0   
364        365   True     40         Resign  Black           10+0   
...        ...    ...    ...            ...    ...            ...   
20013    20014   True     95         Resign  Black          45+45   
20014    20015   True    109           Draw   Draw          45+45   
20015    20016   True     42         Resign  Black           7+30   
20016    20017   True    102         Resign  Black           7+30   
20017    20018   True     97         Resign  White          60+30   

                 white_id  white_rating        black_id  black_rating  \
7      daniel_likes_chess          1413        soultego          2108   
78               oldpaths       
```

<!-- grade_id: cell-5c73407a84990420, points: 1, grade: true, solution: false -->
```python
# Question 3 Grading Checks

assert top_10_percent_black.shape == (2011, 17), 'Make sure that you are selecting the top 10% of black players by rating. Hint: Try using a conditional statement to check which black_rating values are in the top 10%.'

```