# Activity: Using Pandas

## Introduction

In this activity you will practice using some of the basic functionality associated with Pandas.
This activity will cover the following topics:
- Creating a `DataFrame`
- Displaying `DataFrame` information
- Accessing column data
- Getting ranges of column data
- Creating Series from dictionaries
- Using the `iloc()` and `loc()` methods
- Getting data from multiple columns
- Getting rows


#### Question 1

Create a DataFrame called `df` from the given CSV file `student_data.csv` then using the `df` DataFrame, assign the `Name` column to a `Series` called `names`.


<!-- grade_id: cell-98375e57b88587bc, points: 0, grade: false, solution: true -->
```python
import pandas as pd
# Your code here
df = pd.read_csv('student_data-1.csv')
# your code here
names = df['Name']
names
```

Output:
```
0     Jennifer Jackson
1      Michael Johnson
2           Robert Lee
3         Linda Harris
4        Michael Moore
            ...       
70     William Johnson
71         Linda Smith
72        Jennifer Lee
73      Michael Miller
74       Emily Johnson
Name: Name, Length: 75, dtype: object
```

<!-- grade_id: cell-d179546d994dce5f, points: 1, grade: true, solution: false -->
```python
# Question 1 Grading Checks

assert isinstance(df, pd.DataFrame), "Did you create a DataFrame called df?"
assert isinstance(names, pd.Series), "Did you assign the Name column in a variable called names?"

```

#### Question 2

Using the `df` DataFrame, assign the `Age` and `Math Grade` columns to a DataFrame called `age_math`.


<!-- grade_id: cell-0625d61ebce59851, points: 0, grade: false, solution: true -->
```python
# Your code here
age_math = df[['Age','Math Grade']]
# your code here
age_math
```

Output:
```
Age  Math Grade
0    14          84
1    14          92
2    18          87
3    13          90
4    18          88
..  ...         ...
70   17          99
71   18          92
72   17          75
73   12          86
74   13          94

[75 rows x 2 columns]
```

<!-- grade_id: cell-54d4a96060947870, points: 1, grade: true, solution: false -->
```python
# Question 2 Grading Checks

assert isinstance(age_math, pd.DataFrame), "Did you assign the Age and Math Grade columns to a variable called age_math?"

```

#### Question 3

Using the `.loc()` method, assign the `Age` and `Math Grade` columns for the first 30 rows of `df` to a variable called `first_thirty_loc`.


<!-- grade_id: cell-89956eaaf3a9c1a6, points: 0, grade: false, solution: true -->
```python
# Your code here
first_thirty_loc = df.loc[:29 , ['Age','Math Grade']]
# your code here
first_thirty_loc
```

Output:
```
Age  Math Grade
0    14          84
1    14          92
2    18          87
3    13          90
4    18          88
5    12          96
6    18          95
7    16          85
8    18          91
9    18          84
10   17          91
11   17          79
12   17          75
13   16          93
14   15          89
15   13         100
16   14          91
17   13          99
18   12          85
19   14          95
20   18          92
21   15          89
22   12          84
23   16          96
24   12          86
25   17          94
26   15          88
27   12          80
28   13          82
29   17          85
```

<!-- grade_id: cell-618e7ac032208e13, points: 1, grade: true, solution: false -->
```python
# Question 3 Grading Checks

assert isinstance(first_thirty_loc, pd.DataFrame), "Did you correctly assign the first 30 rows to a variable called first_thirty_loc?"

```

#### Question 4

Get the even numbered rows from the `Name` and `English Grade` columns and assign the result to a variable called `even_rows_english`.


<!-- grade_id: cell-e972ac5587c71643, points: 0, grade: false, solution: true -->
```python
# Your code here
even_rows_english = df.iloc[::2] [['Name', 'English Grade']]
# your code here
even_rows_english
```

Output:
```
Name  English Grade
0      Jennifer Jackson             81
2            Robert Lee             80
4         Michael Moore             99
6           Sarah Smith             90
8         Linda Jackson             79
10          Jane Thomas             77
12           Jane Brown             89
14       Michael Wilson             86
16           Mary Jones             76
18      Jennifer Miller             85
20         Susan Wilson             97
22          Emily Moore             81
24           John Moore            100
26          Jane Miller             98
28           John Moore             87
30          Michael Lee             77
32     Jennifer Johnson             86
34         Susan Miller             88
36      William Jackson             93
38        Michael Davis             84
40         Sarah Harris             99
42    Christopher Moore             83
44      Christopher Lee             91
46             John Lee             81
48        Michael Davis             93
50   
```

<!-- grade_id: cell-c1342bc37cb13e08, points: 1, grade: true, solution: false -->
```python
# Question 4 Grading Checks

assert isinstance(even_rows_english, pd.DataFrame), "Did you correctly assign even numbered rows to a variable called even_rows_english? Hint: the first row is index 0."

```

```python

```